## Problem 3 - Applying the classifier (4 points)

Finally, we can bring together the pieces that we have created in problems 1 and 2. Here, your aim is to take
advantage of your new functions and sort a dataset of temperatures in Fahrenheit into four different classes.

### Tips for completing this problem

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-4-problem-3.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail.

### Scores for this problem

**Your score on this problem will be based on following criteria:**

- Creating a new **script** file called `temp_functions.py` with the functions that you wrote in Problems 1 and 2
- Including a docstring in the script file
- Converting and classifying the list of temperatures in Part 2
- Storing the number of temperatures in each class in the specified variables
- Answering some questions at the end of this problem
- Uploading your notebook to your GitHub repository for this week's exercise

## AI tool usage agreement

**Enter your name (and that of your partner) in the cell below** to confirm that you have followed the [course guidelines on the use of AI tools](https://geo-python-site.readthedocs.io/en/latest/course-info/ai-tools.html) and understand that misuse of AI tools is considered cheating.

<!-- grade_id: E4-P3-AI-LLM_statement, points: 0, grade: true, solution: true -->
Матвеев Николай

### Part 1: Creating a script file (1 point)

For this part you should

- Create a new **script file** called `temp_functions.py` in your `exercise-4-username` directory as we saw in the [Lesson 4 materials](https://geo-python-site.readthedocs.io/en/latest/notebooks/L4/script-files.html). Copy and paste the functions that you wrote in Problems 1 and 2 into that script file (i.e., the `fahr_to_celsius` and `temp_classifier` functions)
    - You **do not** need to copy the `ai_temp_classifier` function if you created one in Problem 2
- Add a docstring to the script file (in addition to the separate docstrings in the functions)
- Save the script file and continue working in **this** Notebook (i.e., `Exercise-4-problem-3.ipynb`)

### Part 2: Convert Fahrenheit temperatures to Celsius (3 points)

For this part you should

1. Store the Fahrenheit temperatures below into the list `temp_data` (just execute the existing code cell below).

```python
# List of half-hourly temperature values (in degrees Fahrenheit) for one week
temp_data =  [19, 21, 21, 21, 23, 23, 23, 21, 19, 21, 19, 21, 23, 27, 27, 28, 30, 30, 32, 32, 32, 32, 
              34, 34, 34, 36, 36, 36, 36, 36, 36, 34, 34, 34, 34, 34, 34, 32, 30, 30, 30, 28, 28, 27,
              27, 27, 23, 23, 21, 21, 21, 19, 19, 19, 18, 18, 21, 27, 28, 30, 32, 34, 36, 37, 37, 37, 
              39, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 39, 39, 37, 37, 36, 36, 34, 34, 32, 30, 30,
              28, 27, 27, 25, 23, 23, 21, 21, 19, 19, 19, 18, 18, 18, 21, 25, 27, 28, 34, 34, 41, 37, 
              37, 39, 39, 39, 39, 41, 41, 39, 39, 39, 39, 39, 41, 39, 39, 39, 37, 36, 34, 32, 28, 28,
              27, 25, 25, 25, 23, 23, 23, 23, 21, 21, 21, 21, 19, 21, 19, 21, 21, 19, 21, 27, 28, 32,
              36, 36, 37, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 41, 41, 41, 41, 39, 37, 36, 36, 34,
              32, 30, 28, 28, 27, 27, 25, 25, 23, 23, 23, 21, 21, 21, 19, 19, 19, 19, 19, 19, 21, 23,
              23, 23, 25, 27, 30, 36, 37, 37, 39, 39, 41, 41, 41, 39, 39, 41, 43, 43, 43, 43, 43, 43,
              43, 43, 43, 39, 37, 37, 37, 36, 36, 36, 36, 34, 32, 32, 32, 32, 30, 30, 28, 28, 28, 27,
              27, 27, 27, 25, 27, 27, 27, 28, 28, 28, 30, 32, 32, 32, 34, 34, 36, 36, 36, 37, 37, 37,
              37, 37, 37, 37, 37, 37, 36, 34, 30, 30, 27, 27, 25, 25, 23, 21, 21, 21, 21, 19, 19, 19,
              19, 19, 18, 18, 18, 18, 18, 19, 23, 27, 30, 32, 32, 32, 32, 32, 32, 34, 34, 34, 34, 34,
              36, 36, 36, 36, 36, 32, 32, 32, 32, 32, 32, 32, 32, 30, 30, 30, 30, 30, 30, 30, 30, 30,
              30, 30, 30, 30, 28, 28]
```

#### Part 2 (continues)

2. Import the `fahr_to_celsius` and `temp_classifier` functions from your `temp_functions.py` script
3. Create an empty list called `temp_classes` (which will be filled with temperature class numbers later)
4. **Convert the Fahrenheit temperatures** in the `temp_data` list into Celsius and **classify the Celsius temperatures** using the importend functions. The output should be stored in the `temp_classes` list.
 
    **You can do the temperature conversion and classification inside one for loop:**

    - Iterate over the list of Fahrenheit temperatures (i.e., create a `for` loop), and inside the loop:

        1. Create a new variable called `temp_celsius` in which you assign the temperature in Celsius using the `fahr_to_celsius` function to convert the Fahrenheit temperature into Celsius
        2. Create a new variable called `temp_class` in which you assign the temperature class number (0, 1, 2, or 3) using the `temp_classifier` function
        3. Add the `temp_class` value to the `temp_classes` list

<!-- grade_id: problem_3_ag_p2_functions_usage, points: 0, grade: false, solution: true -->
```python
from temp_functions import *
temp_celsius = []
temp_classes = [] 
for i in temp_data:
    temp_celsius.append(fahr_to_celsius(i))
print (temp_celsius)
for j in range(len(temp_celsius)):
    temp_class = []
    temp_class.append(temp_classifier(temp_celsius[j]))
    temp_class.append(temp_celsius[j])
    temp_classes.append(temp_class)
print (temp_classes)
# Я не до конца понял 3 пункт, поэтому добавил список в список,
# аналогично тому решению, которое описано в задании-3
```

Output:
```
[-7.222222222222222, -6.111111111111111, -6.111111111111111, -6.111111111111111, -5.0, -5.0, -5.0, -6.111111111111111, -7.222222222222222, -6.111111111111111, -7.222222222222222, -6.111111111111111, -5.0, -2.7777777777777777, -2.7777777777777777, -2.2222222222222223, -1.1111111111111112, -1.1111111111111112, 0.0, 0.0, 0.0, 0.0, 1.1111111111111112, 1.1111111111111112, 1.1111111111111112, 2.2222222222222223, 2.2222222222222223, 2.2222222222222223, 2.2222222222222223, 2.2222222222222223, 2.2222222222222223, 1.1111111111111112, 1.1111111111111112, 1.1111111111111112, 1.1111111111111112, 1.1111111111111112, 1.1111111111111112, 0.0, -1.1111111111111112, -1.1111111111111112, -1.1111111111111112, -2.2222222222222223, -2.2222222222222223, -2.7777777777777777, -2.7777777777777777, -2.7777777777777777, -5.0, -5.0, -6.111111111111111, -6.111111111111111, -6.111111111111111, -7.222222222222222, -7.222222222222222, -7.222222222222222, -7.777777777777778, -7.777777777777778, -6.111111111111111, -2.77
```

#### Part 2 (continues)

6. Calculate how many temperatures there are in each temperature class:

    1. Create four variables called `zeros`, `ones`, `twos`, and `threes` 
    2. Count and assign to each variable how many times values 0, 1, 2, and 3 are present in the `temp_classes` list and print out the results below. In other words, assign the value of how many time `0` is in the list to the variable `zeros`, and so on. 
  
**TIP**: You might want to consider using the [**count()** function](https://www.tutorialspoint.com/python3/list_count.htm) OR a for loop for this.

```python
zeros = 0 
ones = 0
twos = 0 
threes = 0
for k in range(len(temp_classes)):
    if (temp_classes[k])[0] == 0:
        zeros += 1
    elif (temp_classes[k])[0] == 1:
        ones += 1
    elif (temp_classes[k])[0] == 2:
        twos += 1
    else: (temp_classes[k])[0] == 3
    threes += 1
# 1. How many 0 values exist in the temp_classes list?
print(zeros)

# 2. How many 1 values exist in the temp_classes list?
print(ones)

# 3. How many 2 values exist in the temp_classes list?
print(twos)

# 4. How many 3 values exist in the temp_classes list?
print(threes)

```

Output:
```
137
85
114
336
```

### Problem 3 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_3_ag_p2_functions_usage_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Check that functions are in the namespace
assert inspect.isfunction(fahr_to_celsius)
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that variable has been created
assert 'temp_celsius' in locals()
assert 'temp_class' in locals()
assert 'temp_classes' in locals()

# Check that temp_classes is a list
assert type(temp_classes) == list
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test3, points: 0.5, grade: true, solution: false -->
```python
# Check that the functions have a single parameter
t_params = list(inspect.signature(temp_classifier).parameters.keys())
f_params = list(inspect.signature(fahr_to_celsius).parameters.keys())
assert len(t_params) == 1
assert len(f_params) == 1
```

Check that required variables exists and print their value (check manually that the answers make sense!):

<!-- grade_id: problem_3_ag_p2_functions_usage_test4, points: 0.5, grade: true, solution: false -->
```python
# Check the variable "zeros" 
assert 'zeros' in locals()
print(zeros)
```

Output:
```
137
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test5, points: 0.5, grade: true, solution: false -->
```python
# Check the variable "ones" 
assert 'ones' in locals()
print(ones)
```

Output:
```
85
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test6, points: 0.5, grade: true, solution: false -->
```python
# Checkthe variable "twos" 
assert 'twos' in locals()
print(twos)
```

Output:
```
114
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test7, points: 0.5, grade: true, solution: false -->
```python
# Check the variable "threes" 
assert 'threes' in locals()
print(threes)
```

Output:
```
336
```

### Part 3: Check your code (0.5 points)
 
- Запомните, чтобы проверить, что у вас были комментарии в ваш код и использовали требуемые имена переменных
- Убедитесь, что ** совершить и подтолкнуть ** свой сценарий `temp_functions.py` в свой собственный репозиторий GitHub для упражнения 4!
- Если у вас есть какие -либо мысли и проблемы на этом этапе, вы можете написать их ниже.

<!-- grade_id: problem_3_mg_p3_check_your_code, points: 0.5, grade: true, solution: true -->
## <span style="color:blue">Применение класса weather из прошлого задания.</span>

Тут я просто показал, как можно сразу разбить нужные данные на "красивый" массив. Не могу судить о хорошей оптимизации (в том смысле, что была бы эта задача с time range, я бы скорее всего точно не использовал класс)

Несмотря на то, что задание предполагает import функции из другого файла, я решил сразу указать
тот же класс и отредактированные методы прям здесь (чтобы вы не переключались между файлами).

```python
class weather():
    def __init__(self, temp_celsius, type_of_weather=None, value = None): 
        """Класс для описания погоды по температуре"""
        self.temp_celsius = temp_celsius
        self.type_of_weather = type_of_weather
        self.value = value

    def classifier(self):
        """Присваивание индекса и типа погоды, в зависимости от температуры"""
        if (self.temp_celsius) < -2: # Как и написано сверху - абсолютная копия из ex.3
            #Однако добавил две строчки снизу
            self.type_of_weather = "cold"
            self.value = 0
        elif (self.temp_celsius) >= -2 and self.temp_celsius<2:
            self.type_of_weather = "slippery"
            self.value = 1
        elif self.temp_celsius  >= 2 and self.temp_celsius<15:
            self.type_of_weather = "comfortable"
            self.value = 2
        elif self.temp_celsius >= 15:
            self.type_of_weather = "warm"
            self.value = 3
def temp_classifier(temp_celsius): # Просто присваиваю число к классу, выполняю метод и "вытаскиваю" только индекс из класса
    temp_celsius = weather(temp_celsius) #Указал класс
    temp_celsius.classifier() #Выполнил метод
    return [temp_celsius.value, temp_celsius.type_of_weather, temp_celsius.temp_celsius]
#Прошу вернуть список в порядке: температура, тип погоды, индекс типа погоды

temp_data =  [19, 21, 21, 21, 23, 23, 23, 21, 19, 21, 19, 21, 23, 27, 27, 28, 30, 30, 32, 32, 32, 32, 
              34, 34, 34, 36, 36, 36, 36, 36, 36, 34, 34, 34, 34, 34, 34, 32, 30, 30, 30, 28, 28, 27,
              27, 27, 23, 23, 21, 21, 21, 19, 19, 19, 18, 18, 21, 27, 28, 30, 32, 34, 36, 37, 37, 37, 
              39, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 39, 39, 37, 37, 36, 36, 34, 34, 32, 30, 30,
              28, 27, 27, 25, 23, 23, 21, 21, 19, 19, 19, 18, 18, 18, 21, 25, 27, 28, 34, 34, 41, 37, 
              37, 39, 39, 39, 39, 41, 41, 39, 39, 39, 39, 39, 41, 39, 39, 39, 37, 36, 34, 32, 28, 28,
              27, 25, 25, 25, 23, 23, 23, 23, 21, 21, 21, 21, 19, 21, 19, 21, 21, 19, 21, 27, 28, 32,
              36, 36, 37, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 41, 41, 41, 41, 39, 37, 36, 36, 34,
              32, 30, 28, 28, 27, 27, 25, 25, 23, 23, 23, 21, 21, 21, 19, 19, 19, 19, 19, 19, 21, 23,
              23, 23, 25, 27, 30, 36, 37, 37, 39, 39, 41, 41, 41, 39, 39, 41, 43, 43, 43, 43, 43, 43,
              43, 43, 43, 39, 37, 37, 37, 36, 36, 36, 36, 34, 32, 32, 32, 32, 30, 30, 28, 28, 28, 27,
              27, 27, 27, 25, 27, 27, 27, 28, 28, 28, 30, 32, 32, 32, 34, 34, 36, 36, 36, 37, 37, 37,
              37, 37, 37, 37, 37, 37, 36, 34, 30, 30, 27, 27, 25, 25, 23, 21, 21, 21, 21, 19, 19, 19,
              19, 19, 18, 18, 18, 18, 18, 19, 23, 27, 30, 32, 32, 32, 32, 32, 32, 34, 34, 34, 34, 34,
              36, 36, 36, 36, 36, 32, 32, 32, 32, 32, 32, 32, 32, 30, 30, 30, 30, 30, 30, 30, 30, 30,
              30, 30, 30, 30, 28, 28]
temp_celsius = []
temp_classes = []
for i in temp_data:
    temp_celsius.append(fahr_to_celsius(i)) # в этом цикле особого смысла нет,
    # можно обойтись без него, но, тк я копировал, он остался
for j in range(len(temp_celsius)):
    temp_classes.append(temp_classifier(temp_celsius[j]))
print (temp_classes)
```

Output:
```
[[0, 'cold', -7.222222222222222], [0, 'cold', -6.111111111111111], [0, 'cold', -6.111111111111111], [0, 'cold', -6.111111111111111], [0, 'cold', -5.0], [0, 'cold', -5.0], [0, 'cold', -5.0], [0, 'cold', -6.111111111111111], [0, 'cold', -7.222222222222222], [0, 'cold', -6.111111111111111], [0, 'cold', -7.222222222222222], [0, 'cold', -6.111111111111111], [0, 'cold', -5.0], [0, 'cold', -2.7777777777777777], [0, 'cold', -2.7777777777777777], [0, 'cold', -2.2222222222222223], [1, 'slippery', -1.1111111111111112], [1, 'slippery', -1.1111111111111112], [1, 'slippery', 0.0], [1, 'slippery', 0.0], [1, 'slippery', 0.0], [1, 'slippery', 0.0], [1, 'slippery', 1.1111111111111112], [1, 'slippery', 1.1111111111111112], [1, 'slippery', 1.1111111111111112], [2, 'comfortable', 2.2222222222222223], [2, 'comfortable', 2.2222222222222223], [2, 'comfortable', 2.2222222222222223], [2, 'comfortable', 2.2222222222222223], [2, 'comfortable', 2.2222222222222223], [2, 'comfortable', 2.2222222222222223], [1, 'slip
```

#### <span style="color:green">Вопрос по Time Range</span>
Каждую пару хочу задать один и тот же вопрос, но по разным причинам откладываю на потом.Вопрос в том, как на различных площадках по написанию кода укладываться во временной промежуток, тк это обычная проблема для питона. Я примерно понимаю, почему она возникает и уже знаком с библиотеками, которые "переводят" код в C-подобный, чтобы значительно его ускорить. Чаще всего, если я не укладываюсь в time range я переписываю код на C++ и все прекрасно работает, но хотелось бы узнать про какой-нибудь способ, чтобы избежать этих переходов с язык на язык (хотя ответ очевиден - перейти на C++ и не жаловаться, что что-то работает долго).

В итоге, прошу Вас поделиться либо литературой, либо объяснить подробнее почему питон так долго выполняет задачи (например цикл в цикле), как лучше всего это избегать (чтобы другой программист потом не говорил какой я дурак).
            Заранее спасибо большое.